<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Einführung Machine Learning
### Sommersemester 2022
Prof. Dr. Heiner Giefers

## Aufgabe: Matrizenmultiplikation in Python

Wir importieren NumPy und erzeugen zwei Matrizen der Größe $N\times{}N$ mit gleichverteilten Zufallszahlen aus dem Bereich $[0,1)$.

```python
import numpy as np

N = 100
A = np.random.rand(N, N)
B = np.random.rand(N, N)
```

Wir wollen nun die beiden NumPy Matrizen miteinander multiplizieren. Dazu können wir den Matrizenmultiplikations-Operator `@` von NumPy verwenden.

Um die Laufzeit dieser Operation zu messen, verwenden wir eine sogenannte *IPython Magic Function* namens `timeit`.
*Magic Functions* werden mit einem vorangestellten `%` aufgerufen:

```python
%timeit C=A@B
```

Nun wollen wir die gleiche Operation in Python, also ohne NumPy implementieren.
Dazu wandeln wir zunächst die *ndarrays* `A` und `B` in *normale* Python Listen-Objekte um:

```python
A_list = A.tolist()
B_list = B.tolist()
type(A_list), type(B_list)
```

Nun sind Sie an der Reihe. Schreiben Sie eine Funktion `matmul` mit der Sie eine Matrizenmultiplikation zweier Matrizen `A` und `B` durchführen. Gehen Sie davon aus, dass `A` und `B` als zweidimensionale Listen (*Listen von Listen*) übergeben werden. Die Funktion soll eine neue Matrix `C` zurückliefern.

![](https://github.com/fh-swf-hgi/intro-ml/raw/main/Praktika/A03/matmult.png)

<!-- grade_id: cell-80d69cb24a5a5248, points: 0, grade: false, solution: true -->
```python
def matmul(A,B):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fbf61c5254ede6bd, points: 0, grade: true, solution: false -->
```python
A_ = [[1,2,3,4],[5,6,7,8],[9,10,11,12]]
B_ = [[1,2,3],[4,5,6],[7,8,9],[10,11,12]]
assert matmul(A_,B_) == [[70, 80, 90], [158, 184, 210], [246, 288, 330]]
```

Testen Sie nun die Laufzeit der selbst implementierten Matrizenmultiplikation.

<!-- grade_id: cell-ef0d2a6d29e95723, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-50aa90906e96fa85, points: 0, grade: true, solution: false -->
```python
assert 'timeit' in _i
```

Wiederholen Sie die Messungen für verschiedene Größen von N.
Sammeln Sie die Luafzeitergebnisse in zwei Listen `run_numpy` und `run_py`.

Um ein Ergebnis von der `timeit` Funktion zu bekommen, rufen Sie die Funktion folgendermaßen auf:

```python
t = %timeit -o C=A@B
print(f"Mittlere Laufzeit = {t.average}s")
```

<!-- grade_id: cell-7de9c51f01cd0f8c, points: 0, grade: false, solution: true -->
```python

run_numpy = []
run_py = []

N_values = [10, 20, 30, 40, 50, 100]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-10b3887c948c65d9, points: 0, grade: true, solution: false -->
```python
assert len(run_numpy)>0
assert len(run_numpy)==len(run_py)
```

Verwenden Sie *Matplotlib*, um beiden Messreihen in einen Graphen zu plotten.

<!-- grade_id: cell-2ea5d945d0df26ee, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8e09e028df381fa6, points: 0, grade: true, solution: false -->
```python
assert 'plt.' in _i
```